# Exercise sheet 5
**Hello everyone!**

**Points: 17 (2 bonus)**

Topics of this exercise sheet are:
* Working with probability distributions

Please let us know if you have questions or problems! <br>
Contact us during the exercise session, on [ADAM](https://adam.unibas.ch/goto_adam_crs_1266890.html), [Piazza](https://piazza.com/class/kzy15kp8s5t6ku), or [via email](https://sada.dmi.unibas.ch/en/teaching/pids22).

Please submit this exercise sheet on **ADAM**.
Naming conventions:

1. Create a folder with your unique identification (usually the first 6 letters of your last name and then 4 numbers). It should be same for all exercise sheet. If you changed, please use the one used to sheet 3.
2. Create a folder inside the previous folder, with name "exerciseX" where X is the sheet number.
3. Place the exercise sheet inside this folder. No dataset, no previous sheet, no remaining of this sheet.
4. Complete the sheet. Only put code or text inside the blocks where "# YOUR CODE HERE" or "YOUR ANSWER HERE" is written. Everything else will be deleted during grading. Don't add new blocks.
5. zip the folder with you name and submit it to the corresponding assignment. No renaming of the zip file either.

Common mistakes:
- Don't use capital letter for your identification name, or the exercise folder.
- Don't put previous sheet or dataset on the submission folders. Just the sheet, ipynb format.



**Handout date**: 2022/04/05 <br>
**Submission date**: 2022/04/12 <br>
**Discussion date**: 2022/04/17

## Introduction:
In this exercise, you get familiar with the basic concepts of probability and statistics. We recommend you to read this article to ensure you know the basic concepts well:
https://towardsdatascience.com/an-in-depth-crash-course-on-random-variables-a3905d03e322

```python
import pandas as pd
import seaborn as sns
import matplotlib.pyplot as plt
import numpy as np
from scipy import stats
from nose.tools import assert_is_instance, assert_equal, assert_almost_equal, assert_true
```

# 1 (6 points)
## Binomial distribution:
The binomial distribution with parameters n and p is a class of the discrete probability distributions of <b> the number of successes in a sequence of $n$ independent experiments, </b> each asking a yes–no question, and each with its own Boolean-valued outcome: success (with probability $p$) or failure (with probability $1 − p$). The probability of happening $x$ times success in $n$ trials is:
$$
    p(x) = \binom{n}{x}p^x (1-p)^{n-x} 
$$

```python
n = 20 # The total number of trials
p = 0.5 # The probability of success in each trial
X = stats.binom(n, p)
# Please check this page for more information: 
# https://docs.scipy.org/doc/scipy-0.13.0/reference/generated/scipy.stats.binom.html#scipy.stats.binom
```

### 1a (2 points)
Generate 10000 random samples from the given binomial distribution and plot the histogram.

Hint: Use X.rvs()

<!-- grade_id: cell-cda45808f0efe171, points: 0, grade: false, solution: true -->
```python
n_samples=10000
np.random.seed(0)
# YOUR CODE HERE
raise NotImplementedError()
# x = 
```

<!-- grade_id: cell-bf5d9fe64b7e6143, points: 2, grade: true, solution: false -->
```python
print(x.mean())
print(x.std())
assert_equal(len(x), n_samples)
assert_almost_equal(x.mean(), 9.968 , places=2)
assert_almost_equal(x.std() , 2.249 , places=2)
```

### 1b (2 points)
What are $\text{P}(X = 8)$, $\text{P}(X = 10)$ and $\text{P}(X = 12)$?

Hint: Use X.pmf()

<!-- grade_id: cell-183fea12257ba84f, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
# p8 =
# p10 =
# p12 =
```

<!-- grade_id: cell-b16bddc2f8788742, points: 2, grade: true, solution: false -->
```python
print((p8 - p12)/p10)
print((p10*p12)+p8)

assert_almost_equal((p8 - p12)/p10, 0)
assert_almost_equal((p10*p12)+p8, 0.141 , places=2)
```

Which point has the highest probability? Why?

<!-- grade_id: cell-b81018d863cac204, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

### 1c (1 points bonus)
What are $\text{P}(X \leq 8)$, $\text{P}(X \leq 10)$, $\text{P}(X \leq 12)$ and $\text{P}(X \leq 20)$?

Hint: Use X.cdf()

<!-- grade_id: cell-3267fe45f8b26292, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
# p_le8 = 
# p_le10 = 
# p_le12 = 
# p_le20 = 
print('P(X<=8)={:.4f}\nP(X<=10)={:.4f}\nP(X<=12)={:.4f}\nP(X<=20)={:.4f}'.format(p_le8, p_le10, p_le12,p_le20))

```

<!-- grade_id: cell-0ed57cd19a0323c4, points: 1, grade: true, solution: false -->
```python
print(p_le10 - p_le8)
print((p_le12 - p_le20)/p_le8)

assert_almost_equal(p_le10 - p_le8, 0.336 , places=2)
assert_almost_equal((p_le12 - p_le20)/p_le8 , -0.522 , places=2)
```

Do you see an increasing trend? Is $\text{P}(X \leq 20)$ the highest value and equals 1? Why? 

<!-- grade_id: cell-7cd7e3ce262842e4, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

### 1d (2 points)
What are the mean, standard deviation and median of the generated samples?

Hint: You can use numpy!

<!-- grade_id: cell-0dface407979af48, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

print((mean - std)/median)
print((mean ** std)+ median)
```

<!-- grade_id: cell-554077d683ce6e56, points: 2, grade: true, solution: false -->
```python
assert_almost_equal((mean - std)/median, 0.771 , places=2)
assert_almost_equal((mean ** std)+ median, 186.349 , places=2)
```

Is the mean close to $n \times p$? Why should we expect this?

<!-- grade_id: cell-c213f4d54ab18ec1, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

# 2 (6 points)
## Multimodal distributions:
Now we consider samples from two binomial distributions with different p values.

```python
n = 20
p1 = 0.2
p2 = 0.8
X1 = stats.binom(n, p1)
X2 = stats.binom(n, p2)
```

### 2a (2 points)

Generate 5000 random samples from each of these two binomial distributions and concatenate the samples to make a unified vector. Plot the histogram of the concatenated samples.

<!-- grade_id: cell-eb08c2c0ae675c5a, points: 0, grade: false, solution: true -->
```python
n_samples=5000
np.random.seed(0)

# YOUR CODE HERE
raise NotImplementedError()
# x1 = 
# x2 = 
# x = 
```

<!-- grade_id: cell-80410be6ed2be347, points: 2, grade: true, solution: false -->
```python
print((x1.mean() + x2.mean())/x.mean())
print((x1.mean() - x2.mean())*x.mean())

assert_almost_equal((x1.mean() + x2.mean())/x.mean(), 2.000 , places=2)
assert_almost_equal((x1.mean() - x2.mean())*x.mean(), -120.670 , places=2)
```

### 2b (4 points)

Calculate the mean, median and the first two modes of the concatenated samples.

Hint: You can use np.unique and np.argsort to find the first and the second modes

<!-- grade_id: cell-fd06b5e4a5b3d882, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
# mean = 
# median = 

# mode1 = 
# mode2 = 

print('Mean: {:.2f} \nMedian: {:.2f} \nFirst mode: {:.2f}\nSecond mode: {:.2f}'.format(mean, median, mode1, mode2))
```

<!-- grade_id: cell-f03ecd27589d8c8b, points: 4, grade: true, solution: false -->
```python
print((mean - median)/mode1)
print((mode1 - mode2)/ mean)

assert_almost_equal((mean - median)/mode1, 0.000 , places=2)
assert_almost_equal((mode1 - mode2)/ mean, 1.198 , places=2)
```

Which of these is a better represetation of this type of data? Why?

<!-- grade_id: cell-f441612acd70e110, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

# 3 (3 points)

## Normal distribution:
Normal (Gaussian) distribution is a class of continuous probability distributions with mean $\mu$ and standard deviation $\sigma$:
$$
    p(x) = \frac{1}{\sqrt{2\pi\sigma^2}} e^{-\frac{(x - \mu)^2}{2\sigma^2}}
$$

```python
mu = 3.0
sigma = 2.0
X = stats.norm(loc = mu , scale = sigma)

# Please check this page for more information: 
# https://docs.scipy.org/doc/scipy/reference/generated/scipy.stats.norm.html
```

### 3a (1 points)

Generate 10000 random samples from the the given normal distribution and plot the histogram.

<!-- grade_id: cell-5c2f4c1e4bcf3221, points: 0, grade: false, solution: true -->
```python
n_samples=10000
np.random.seed(0)

# YOUR CODE HERE
raise NotImplementedError()
# x = 
```

<!-- grade_id: cell-9e8e1c6162c85b13, points: 1, grade: true, solution: false -->
```python
print(x.mean())
print(x.std())

assert_equal(len(x), n_samples)
assert_almost_equal(x.mean(), 2.963 , places=2)
assert_almost_equal(x.std() , 1.975 , places=2)
```

### 3b (1 points)

What are $\text{P}(X = 2)$, $\text{P}(X = 3)$ and $\text{P}(X = 4)$?

<!-- grade_id: cell-3f38274cf814ddd3, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
# p2 = 
# p3 = 
# p4 = 
print('P(X=2)={:.4f}|P(X=3)={:.4f}|P(X=4)={:.4f}'.format(p2, p3, p4))
```

<!-- grade_id: cell-5c7c63d65fc8333e, points: 1, grade: true, solution: false -->
```python
assert_equal((p2 - p3), 0)
assert_equal((p2*p3)+p4, 0)
```

### 3c (1 points bonus)

What are $\text{P}(X \leq 2)$, $\text{P}(2 < X \leq 3)$ and $\text{P}(X \geq 4)$?

<!-- grade_id: cell-464078c0201e8ad8, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
# p_le2 = 
# p_g2_le3 = 
# p_ge4 = 

print('P(X<=2)={:.4f}\nP(2<X<=3)={:.4f}\nP(X>=4)={:.4f}'.format(p_le2, p_g2_le3, p_ge4))
```

<!-- grade_id: cell-8d24f229bf16b73e, points: 1, grade: true, solution: false -->
```python
print(p_le2 - p_ge4)
print((p_g2_le3 - p_le2)/p_ge4)

assert_equal(p_le2 - p_ge4, 0)
assert_almost_equal((p_g2_le3 - p_le2)/p_ge4 , -0.379 , places=2)
```

Are $\text{P}(X \leq 2)$ and $\text{P}(X \geq 4)$ equal? Why?

<!-- grade_id: cell-dda20287379f7a87, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

### 3d (1 points)

What are the mean, standard deviation, and median of the generated samples?

<!-- grade_id: cell-67e48526cb3d882c, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
# mean = 
# std =
# median = 
print('Mean={:.4f}|Std={:.4f}|Median={:.4f}'.format(mean, std, median))
```

<!-- grade_id: cell-424c1572d1adec05, points: 1, grade: true, solution: false -->
```python
print((mean - std)/median)
print((mean ** std)+ median)

assert_almost_equal((mean - std)/median, 0.335 , places=2)
assert_almost_equal((mean ** std)+ median, 11.491 , places=2)
```

Are the mean and std close to the $\mu$ and $\sigma$? How about the median? Why?

<!-- grade_id: cell-2afc01e0b24fcdd9, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE