# Exercise sheet 10
**Hello everyone!**

**Points: 15 (+2)**

Please let us know if you have questions or problems! <br>
Contact us during the exercise session, on [ADAM](https://adam.unibas.ch/goto_adam_crs_1266890.html), [Piazza](https://piazza.com/class/kzy15kp8s5t6ku), or [via email](https://sada.dmi.unibas.ch/en/teaching/pids22).

Please submit this exercise sheet on **ADAM**.
Naming conventions:

1. Make a folder called "exercise10".
2. Put your submission "Exercise sheet 10.ipynb" in there.
3. Complete the sheet. Only put code or text inside the blocks where "# YOUR CODE HERE" or "YOUR ANSWER HERE" is written. Everything else will be deleted during grading. Don't add new blocks.
4. Then zip the folder called "exercise10". This will create a zip file called "exercise10.zip". Rename that zip file to your Unibas short name, e.g. "blabla0000.zip".

Common mistakes:
- Don't use capital letter for your identification name, or the exercise folder.
- Don't put previous sheet or dataset on the submission folders. Just the sheet, ipynb format.



**Handout date**: 2022/05/17 <br>
**Submission date**: 2022/05/24 <br>

```python
import pandas as pd
import seaborn as sns
import matplotlib.pyplot as plt
import numpy as np
from scipy import stats
from scipy.stats import norm
from nose.tools import assert_is_instance, assert_equal, assert_almost_equal, assert_true
```

<p><a href="https://commons.wikimedia.org/wiki/File:Basel_-_M%C3%BCnsterpfalz3.jpg#/media/Datei:Basel_-_M체nsterpfalz3.jpg"><img src="https://upload.wikimedia.org/wikipedia/commons/1/18/Basel_-_M%C3%BCnsterpfalz3.jpg" alt="Basler M체nster mit Pfalz und der Mittleren Rheinbr체cke" height="486" width="1280"></a></p>

# Basel's neighborhoods
We are working with a dataset containing information about Basel and its neighborhoods. You can find it here: https://opendata.swiss/de/dataset/kennzahlen-zu-den-basler-wohnvierteln-und-landgemeinden. It's also uploaded as a csv to this folder.

```python
bs = pd.read_csv("basel_neighborhoods.csv", sep=";")
bs.head()
```

## P value (12 points)
Let's see if there's an answer to the following question: Did the amount of green spaces in Basel's neighborhoods increase *significantly* between 2015 and 2021? <br>

The column that measures green spaces (parks and such) is called *anteil_gruenflaechen*.

### Observed difference (1 point)
What is the observed difference between the mean of the amount of green spaces in 2015 and in 2021? <br>
(Hint: Averaged over all neighborhoods)

<!-- grade_id: cell-c1d5d57296720457, points: 0, grade: false, solution: true -->
```python
# mean_2015 = ...
# mean_2021 = ...

# YOUR CODE HERE
raise NotImplementedError()

print(mean_2021 - mean_2015)
```

<!-- grade_id: cell-3b18cf85b244c23e, points: 1, grade: true, solution: false -->
```python
assert_almost_equal(mean_2021 - mean_2015, 0.0102, places=2)
```

### Thinking about it...
Do you think that this difference points to a strong increase in green spaces between 2015 and 2021? What do you think the p value will be approximately?

<!-- grade_id: cell-f407cf89d324922e, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

### Artificial variables (3 points)
Let's make a new dataframe "neighborhoods" containing the neighborhoods of Basel in a column called "wohnviertel_name". Additionally, it should have a column called "increase_green_space" that has...
* a 1 if the amount of green space increased between 2015 and 2021 for this neighborhood
* a 0 otherwise

<!-- grade_id: cell-ce997035c2a68011, points: 0, grade: false, solution: true -->
```python
# neighborhoods = ...
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-6352e5648d9d870a, points: 3, grade: true, solution: false -->
```python
assert_is_instance(neighborhoods, pd.DataFrame)
assert_equal(neighborhoods[neighborhoods["wohnviertel_name"] == "Bachletten"]["increase_green_space"].iloc[0], 1)
assert_equal(neighborhoods[neighborhoods["wohnviertel_name"] == "Matth채us"]["increase_green_space"].iloc[0], 0)
```

### Amount of neighborhoods with an increase in green spaces (3 points)
What is the percentage of neighborhoods that had an increase in green spaces from 2015 to 2021? This value will let us calculate the difference to our null hypothesis value.

<!-- grade_id: cell-707efa52cfcacb2c, points: 0, grade: false, solution: true -->
```python
# The percentage of neighborhoods that saw an increase of green spaces
# increase_2021 = ...
# The difference of increase_2021 to our null hypothesis value
# difference_null_hypothesis = ...
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-6c4890aad81d1b54, points: 3, grade: true, solution: false -->
```python
assert_is_instance(increase_2021, float)
assert_is_instance(difference_null_hypothesis, float)
assert_almost_equal(increase_2021, 0.57142857, places=2)
assert_almost_equal(difference_null_hypothesis, 0.0714285, places=2)
```

### Calculate the p value (3 points)
What is the p value for there being a significant difference in the amount of green space between 2015 and 2021? <br>
Or: What is the p value for $P(|mean2021 - 0.5| > \text{observed_difference}) $?

<!-- grade_id: cell-07f7fc5a4f401a11, points: 0, grade: false, solution: true -->
```python
# p_value = ...
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-c1f16ff9645226be, points: 3, grade: true, solution: false -->
```python
assert_is_instance(p_value, float)
```

### Do we reject or fail to reject the null hypothesis based on our p value? (2 points)
(Hint: Read the question carefully and google it if you're not sure what it means)

<!-- grade_id: cell-e09f38db783d9e60, points: 0, grade: false, solution: true -->
```python
# Please write your answer like this:
# If we accept, comment out the next line:
# null_hypothesis = "reject"
# If we reject, comment out the next line:
# null_hypothesis = "fail to reject"
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-ce8f3510a55faf06, points: 2, grade: true, solution: false -->
```python
assert_is_instance(null_hypothesis, str)
assert_true(null_hypothesis in ("fail to reject", "reject"))
```

### How could we get a lower p value?

<!-- grade_id: cell-7f58662189aa532a, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

## Regression (3 points (+2))
### Plot correlation
Make a plot that shows you how income tax (*einkommenssteuer_pro_veranlagung*) and apartment size (*flaeche_pro_wohnung*) are correlated (a scatter plot) for the year 2021.

<!-- grade_id: cell-30fc240c2ed71fb4, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

### Correlation coefficient (1 point)
What is the correlation coefficient between income tax and apartments size in Basel for 2021?

<!-- grade_id: cell-5a5ef833de831dee, points: 0, grade: false, solution: true -->
```python
# corr_coeff = ...
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-917bc90470c47113, points: 1, grade: true, solution: false -->
```python
assert_is_instance(corr_coeff, float)
```

### Highest correlation (+ 2 bonus points)
What is the highest correlation value between two columns in the bs dataset for the year 2021?

<!-- grade_id: cell-b841e41740bd4cfe, points: 0, grade: false, solution: true -->
```python
# highest_correlation = ...
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-9b623c8c3cb9df90, points: 2, grade: true, solution: false -->
```python
assert_is_instance(highest_correlation, float)
```

### Plot another correlation
Make a plot that shows you how living space per person (*wohnflaeche_pro_person*) and apartment size (*flaeche_pro_wohnung*) are correlated (a scatter plot) for the year 2021.

<!-- grade_id: cell-c1967853363f5c89, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

What do you notice on this plot? What kind of value do you expect for the correlation coefficient? What could be problematic about this correlation?

<!-- grade_id: cell-bbd1af1856fa03f5, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

### Expectations (2 points)
What is the expected value $\hat E$ for the amount of living space for a random person (*wohnflaeche_pro_person*) living in Basel in 2021?

<!-- grade_id: cell-540ceba74e68d292, points: 0, grade: false, solution: true -->
```python
# living_space = ...
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-34ce7292b3c81829, points: 2, grade: true, solution: false -->
```python
assert_is_instance(living_space, float)
```

```python

```