Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

# Exercise sheet 3
**Hello everyone!**

This is the second graded exercise sheet. <br>
**Points: 15(+2)**

Topics of this exercise sheet are:
* Get confident with Jupyter, pandas
* Matplotlib
* Introduction to more advanced visualization tools
* Search new functions on your own

Please let us know if you have questions or problems! <br>
Contact us during the exercise session, on [ADAM](https://adam.unibas.ch/goto_adam_crs_1266890.html), [Piazza](https://piazza.com/class/kzy15kp8s5t6ku), or [via email](https://sada.dmi.unibas.ch/en/teaching/pids22).

Please submit this exercise sheet on **ADAM**.
Naming conventions:
* Name of notebook: "Exercise sheet 3.ipynb"
* Please put it in a folder structure like this: "your_unibas_short_name/Exercise3/Exercise sheet 3.ipynb"
* Your Unibas short name should (maybe) have the first 6 letters of your last name and then 4 numbers. Find it on viaweb.unibas.ch
* Then zip this folder and upload it on ADAM.

**Handout date**: 2022/03/21 <br>
**Submission date**: 2022/03/28 <br>
**Discussion date**: 2022/03/29

##  About this notebook

In this notebook we will explore data on college majors sourced from the American Community Survey 2010–2012 Public Use Microdata Sample. This dataset can be found at this [URL](https://raw.githubusercontent.com/fivethirtyeight/data/master/college-majors/recent-grads.csv).

---
## Part A (2 point)

In this part you will load the dataset and have a rapid overview of what it looks like.

```python
import pandas as pd
from nose.tools import assert_is_instance, assert_equal, assert_almost_equal, assert_true
```

Load the data in a Pandas dataframe using "read_csv" Pandas function.

<!-- grade_id: cell-b29708797b6d4e98, points: 0, grade: false, solution: true -->
```python
# df = pd.read_csv()
# YOUR CODE HERE


# YOUR CODE HERE
raise NotImplementedError()
```

Display the first few rows and **all the columns**. This will help you to have an idea of what is inside the dataset.

<!-- grade_id: cell-4d97c593e4c05a0f, points: 0, grade: false, solution: true -->
```python
# df.head()
# YOUR CODE HERE

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-494f50ef1efaef44, points: 2, grade: true, solution: false -->
```python
test = ((len(df.columns)==pd.options.display.max_columns) or (pd.options.display.max_columns == None) or (pd.options.display.max_columns >= 21))
assert_equal(test,True)
```

## Part B (4 points)

In this part we will plot basic information to get a sense of the dataset. You can use the pandas plot functions or matplotlib plot functions. <br>
For your information, "Rank" is the major’s rank by median earnings.

Does the unemployment rate look proportional to the rank of the major? Display the unemployment rate against the rank of the major.

<!-- grade_id: cell-acfaa3f32953bcbb, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE

# YOUR CODE HERE
raise NotImplementedError()
```

Does the salary look proportional to the rank of the major? Display the median salary, the 25th and 75th percentiles against the rank of the mjor. Display all three lines in the same plot.

<!-- grade_id: cell-b5d156584da87cd5, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE

# YOUR CODE HERE
raise NotImplementedError()
```

Display the number of low wage jobs against the rank of the Major.

<!-- grade_id: cell-ddbbd4c867c01c29, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE

# YOUR CODE HERE
raise NotImplementedError()
```

Can we make conclusions about the relation between low wage jobs and the rank of the Major?

We will display the percentage of low wage jobs against the rank of the major and use Matplotlib to get a more precise plot.

```python
import matplotlib.pyplot as plt
```

Extract the rank, number of low wage jobs, and the number of employed from the data and put each of them into their own series.

<!-- grade_id: cell-2d29ee1c10fccd4f, points: 0, grade: false, solution: true -->
```python
# df_rank = ...
# df_wage = ...
# df_employed = ...

# YOUR CODE HERE

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-2fa9411434ff0423, points: 2, grade: true, solution: false -->
```python
assert_equal(type(df_rank), pd.core.series.Series)
assert_equal(df_rank.name, 'Rank')
assert_equal(type(df_wage), pd.core.series.Series)
```

Display the percentage of low wage jobs against the rank of the major. 

<!-- grade_id: cell-c8617292a46c524d, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE

# YOUR CODE HERE
raise NotImplementedError()
```

## Part C (3 point, +1)
In this part, we will get some information about the distribution of the data.

Display the histogram for the median of the salary. 

<!-- grade_id: cell-d6d3efd95686ee59, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE

# YOUR CODE HERE
raise NotImplementedError()
```

Display the histogram for the ratio of employed with low wage. <br>
**Hint**: use matplotlib to do this plot, and the variable 'df_wage', 'df_employed' that you extracted previously.

<!-- grade_id: cell-1caf4c0fa974197f, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE

# YOUR CODE HERE
raise NotImplementedError()
```

Use subplot to display two figures.<br>
The first figure is the histogram for the ratio of employed with low wage (like the previous cell).<br>
The second figure is the same histogram but with 30 equal-width bins.

<!-- grade_id: cell-d36fecee97e180f0, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Make the previous plot pretty. <br>
Add one title by subfigure. Use two different colors for the two histograms. Add x and y labels. Display a vertical grid in the first sub-figure, and a horizontal grid on the second one. Remove the edges around the plot. <br>
We will give a +1 bonus for very very cute plots!

<!-- grade_id: cell-0e0738d9263bbb86, points: 3, grade: true, solution: true -->
```python
# YOUR CODE HERE

# YOUR CODE HERE
raise NotImplementedError()
```

## Part D (2 points)
Based on the histogram that you displayed in the previous part, you should see some isolated extreme values.
In this part, we will try to identify and understand these outliers.

Extract a dataframe containing the rows with the 10 highest median salaries.

<!-- grade_id: cell-053de464e1e1e944, points: 0, grade: false, solution: true -->
```python
# top_10 = ...

# YOUR CODE HERE

# YOUR CODE HERE
raise NotImplementedError()
```

From the Major with top 10 median salary, display the median, the 25th and 75th percentiles on a bar plot.<br>
**Hint**: use the attribute 'kind' in the pandas plot function.

<!-- grade_id: cell-7303c8f3fd12b522, points: 1, grade: true, solution: false -->
```python
assert_equal(type(top_10), pd.core.frame.DataFrame)
```

<!-- grade_id: cell-97fda96bd4b5fa3e, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE

# YOUR CODE HERE
raise NotImplementedError()
```

We can also plot the histogram for a specific category. Display the histogram of the median salary for the Major with category "Engineering".

<!-- grade_id: cell-f89fca6713c753e0, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE

# YOUR CODE HERE
raise NotImplementedError()
```

## Part E
Let's now see if we can infer the correlation between two factors using scatter plots.

Display the median salary against the unemployment rate on a scatter plot.<br>
**Hint**: use once again the attribute 'kind' from the pandas plot function.

<!-- grade_id: cell-952575fb38116382, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE

# YOUR CODE HERE
raise NotImplementedError()
```

Do the same plot using matplotib.<br>
**Hint**: You can use function plt.plot() or plt.scatter()

<!-- grade_id: cell-981687c0b197b302, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE

# YOUR CODE HERE
raise NotImplementedError()
```

## Part F (4 points, +1)
In this part, we will now see how to analyze and visualize categorical data.

We want to count how many students there are per category. <br>
Extract the different categories (column "Major_category") and display them on an **horizontal** bar plot.<br>
Sum up the number of students in each category (column "Total"). <br>
**Hint**: use once again the attribute 'kind' from the pandas plot function.

<!-- grade_id: cell-762e561165e3b73b, points: 0, grade: false, solution: true -->
```python
# cat_totals = ...

# YOUR CODE HERE

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-3a6e2c5cf797488e, points: 2, grade: true, solution: false -->
```python
assert_equal(len(cat_totals),16)
assert_equal(cat_totals["Interdisciplinary"],12296)
assert_equal(cat_totals["Engineering"],537583)
```

Vertical and horizontal bar charts are often a good choice if you want to see the difference between your categories. If you’re interested in ratios, then pie plots are a good tool. However, since cat_totals contains a few smaller categories, creating a pie plot with cat_totals.plot(kind="pie") will produce several tiny slices with overlapping labels.

<!-- grade_id: cell-0c027c0577c89787, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

To address this problem, we will see how to merge the smaller categories into a single group. <br>
First, separate the data in two groups, depending on the number of occurence. The small group should contain all major category that have less than 200,000 individual. <br>
Make sure that the variable containing all the small groups is named "small_cat_totals".

<!-- grade_id: cell-de85dfc4feaed5ab, points: 0, grade: false, solution: true -->
```python
# small_cat_totals = ...
# big_cat_totals = ...

# YOUR CODE HERE

# YOUR CODE HERE
raise NotImplementedError()
```

Merge all small categories into a single one.

<!-- grade_id: cell-0a0cc089c61d3e0c, points: 0, grade: false, solution: true -->
```python
# small_sums = pd.Series(..., index=["Other"])

# YOUR CODE HERE

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-40a3e78a7696fbf1, points: 2, grade: true, solution: false -->
```python
assert_equal(small_cat_totals["Agriculture & Natural Resources"], 75620)
assert_equal(small_cat_totals["Interdisciplinary"], 12296)
```

Add the "other" category to the "big_cat_totals" dataframe. Display the new pie chart. <br>
**Hint**: you can use "pd.concat(...)" to concatenate two elements.

<!-- grade_id: cell-bee4117ea34fb564, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE

# YOUR CODE HERE
raise NotImplementedError()
```

What is the problem of this pie chart? Is the category other well chosen?

<!-- grade_id: cell-a84c420b1aac4032, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

Can you propose a better way to visualize this information (other than with a pie chart)?

<!-- grade_id: cell-b346c4cc72dfe43a, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE