# Exercise sheet 9
**Hello everyone!**

**Points: 15**

Please let us know if you have questions or problems! <br>
Contact us during the exercise session, on [ADAM](https://adam.unibas.ch/goto_adam_crs_1266890.html), [Piazza](https://piazza.com/class/kzy15kp8s5t6ku), or [via email](https://sada.dmi.unibas.ch/en/teaching/pids22).

Please submit this exercise sheet on **ADAM**.
Naming conventions:

1. Make a folder called "exercise9".
2. Put your submission "Exercise sheet 9.ipynb" in there.
3. Complete the sheet. Only put code or text inside the blocks where "# YOUR CODE HERE" or "YOUR ANSWER HERE" is written. Everything else will be deleted during grading. Don't add new blocks.
4. Then zip the folder called "exercise9". This will create a zip file called "exercise9.zip". Rename that zip file to your Unibas short name, e.g. "blabla0000.zip".

Common mistakes:
- Don't use capital letter for your identification name, or the exercise folder.
- Don't put previous sheet or dataset on the submission folders. Just the sheet, ipynb format.



**Handout date**: 2022/05/10 <br>
**Submission date**: 2022/05/17 <br>

```python
import pandas as pd
import seaborn as sns
import matplotlib.pyplot as plt
import numpy as np
from scipy import stats
from scipy.stats import norm
from nose.tools import assert_is_instance, assert_equal, assert_almost_equal, assert_true
```

# More uncertainty and confidence intervals!
In this notebook we will get more and more confidence on confidence intervals. 
This sheet should be handled quickly if you understood the two previous sheets well (especially sheet 7).
This notebook contains more mathematical notations that usually. However, there are here only to make the exercise clear. You won't need to run difficult computations other than empirical mean and empirical standard deviation.


# Titanic dataset
We are working with a dataset containing information about passengers of the sadly famous Titanic. We will see how we can use tools from the lecture to get a sense of what factors are important to increase our probability of survival during a shipwreck.<br>

In this first part, we will load the dataset, and plot a few values to see what could be the important parameters. In a second part, we will focus on one particular feature and see if you can trust the dataset.

We first load the dataset. If you are working locally, you should be able to load the dataset from the csv file name 'train.csv', in the same location as this current file.

```python
import pandas as pd
data_source = "https://raw.githubusercontent.com/paulhendricks/titanic/master/inst/data-raw/train.csv"
data = pd.read_csv(data_source)
data.head()
```

Among the passengeners that survived, display the proportion (histogram) of each 'Pclass'.<br>
On the same plot, display the same histogram for the passengers that didn't survive.<br>
Hint: you can use plt.hist.

<!-- grade_id: cell-ede3d355a0957fbe, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Among the passengeners that survived, display the proportion (histogram) of each 'Fare'.<br>
On the same plot, display the same histogram for the passengers that didn't survive.<br>
Hint: you can use plt.hist.

<!-- grade_id: cell-f95caf2cb1886049, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Among the passengeners that survived, display the proportion (histogram) of each age.<br>
On the same plot, display the same histogram for the passengers that didn't survive.<br>
Hint: you can use plt.hist.

<!-- grade_id: cell-0da4e66b7921cec4, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Display on the same plot, the proportion of female that survived and the proportion of male. Hint: you can use plt.hist.<br>
Among the passengers, compute the proportion of survivors that are women and store it in *surv_female*.<br>
Among the passengers, compute the proportion of survivors that are men and store it in *surv_male*.


<!-- grade_id: cell-a66ee7dc2a8b5ed1, points: 0, grade: false, solution: true -->
```python
surv_female = 0
surv_male = 0 
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-98737fddf6b72220, points: 2, grade: true, solution: false -->
```python
print(surv_female)
print(surv_male)
```

## Will your sex increase or decrease your probability of survival?

Let $(X_i)_{i\leq N_X} \overset{i.i.d.}{\sim} \mathrm{Ber}(p_X)$ denote $N_X$ Bernoulli random variables that model the outcome of a female on the Titanic. <br>
Let $(Y_i)_{i\leq N_Y} \overset{i.i.d.}{\sim} \mathrm{Ber}(p_Y)$ denote $N_Y$ Bernoulli random variables that model the outcome of a male on the Titanic.<br>

We aim at estimating the probailty of survival $p_X$ and $p_Y$.

Ideally we would like to get the values $p_X = \mathbb{E}[X_i]$ and $p_Y = \mathbb{E}[Y_i]$. Unfortunately, we only have access to the $N_X$ realizations $(X_i)$ and $N_Y$ realization $(Y_i)$. <br>

Use the law of large numbers to create an estimate $\hat p_X$ of $p_X$ and an estimate $\hat p_Y$ of $p_Y$. Name these estimators *pX_hat* and *pY_hat*.

<!-- grade_id: cell-f870da6474590bea, points: 0, grade: false, solution: true -->
```python
pX_hat = 0
pY_hat = 0
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-604fcaba3e728967, points: 2, grade: true, solution: false -->
```python
print(pX_hat)
print(pY_hat)
assert_almost_equal(pX_hat , 0.7420382165605095 , places=2)
assert_almost_equal(pY_hat , 0.18890814558058924 , places=2)
```

Let's do some math (or follow slides of lecture 10).<br>
We have some estimate for $p_X$ and $p_Y$, but we would like to know if we can trust them. <br>

Let's first focus of $p_X$. <br>
We want to compute the probability $p_{trust X}$
$$ p_{trust X} = \mathbb{P}\left( p_X-0.05 \leq \hat X \leq p + 0.05 \right). $$

Based the central limit theorem, we'll make the assumption that $\hat p_X \sim \mathcal{N}(p_X, \sigma^2/N_X)$.<br>
Then, following slides of lecture 10 (p18, p19), we have
$$ p_{trust X} = \mathbb{P}\left( Z \leq 0.05\sqrt{ \frac{N_X}{p_X(1-p_X)}} \right)-\mathbb{P}\left( Z \leq -0.05\sqrt{ \frac{N_X}{p_X(1-p_X)}} \right), $$
where $Z\sim \mathcal{N}(0,1)$. <br>
As mentioned during the lecture, we don't know $p_X(1-p_X)$. Let's replace this term by $\hat p_X(1-\hat p_X)$.




Compute the value 
$$ p1 = \mathbb{P}\left( Z \leq 0.05\sqrt{ \frac{N_X}{\hat p_X(1-\hat p_X)}} \right)$$
and 
$$ p2 = \mathbb{P}\left( Z \leq -0.05\sqrt{ \frac{N_X}{\hat p_X(1-\hat p_X)}} \right).$$
Store these values in variables named $p1$ and $p2$.<br>
Hint: you can have access to the cumulative distribvutuion function (cdf) of a Gaussian random variable by using the function *norm* (for normal) form the library *scipy.stats* (already imported).

```python
# Leave this block. 
# This is to make sure that your next answer can still bring you points even if you make a mistake previously.
pX_hat = 0.7420382165605095
pY_hat = 0.18890814558058924
```

<!-- grade_id: cell-11532a187e7a16cc, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-20da3526a3f1f163, points: 3, grade: true, solution: false -->
```python
print(p1, p2)
assert_almost_equal(p1 , 0.978570861003359 , places=2)
assert_almost_equal(p2 , 0.02142913899664103 , places=2)
```

Based on the previous assumptions, what is the probability that $\hat p_X \in [p_X - 0.05, p_X+0.05]$? Store your answer in a variable named *proba_trust_X*.

```python
# Leave this block. 
# This is to make sure that your next answer can still bring you points even if you make a mistake previously.
p1 = 0.978570861003359
p2 = 0.02142913899664103 
```

<!-- grade_id: cell-c5f50be87fabc3a0, points: 0, grade: false, solution: true -->
```python
proba_trust_X = 0
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-84718dfa47b17516, points: 2, grade: true, solution: false -->
```python
print(proba_trust_X)
assert_almost_equal(proba_trust_X , 0.957141722006718 , places=2)
```

Using a similar approach, compute the probability that $\hat p_Y \in [p_Y - 0.05, p_Y+0.05]$? Store your answer in a variable named *proba_trust_Y*.

<!-- grade_id: cell-b6aeee80244a9dd3, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-b361b39476a7faf5, points: 4, grade: true, solution: false -->
```python

```

# What about the children?
Should we remove the children in the previous analysis about the sex? If we do so, the number of women drop from 314 to 259. And we have seen in the lecture that the lower the number of sample, the worse is the confidence interval. 
<br>
But anyway, let's finish by checking that children have a very high probability of survival. Compute the probability of survival of a child (age<18) and store it in *proba_surv_child*.


<!-- grade_id: cell-f85e1d8e3c4f0049, points: 0, grade: false, solution: true -->
```python
proba_surv_child = 0
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-5e346eed49d5396b, points: 2, grade: true, solution: false -->
```python
print(proba_surv_child)
```