# Exercise sheet 6
**Hello everyone!**

**Points: 15**

Topics of this exercise sheet are:
* Probabilities
* Conditional Probabilities

Please let us know if you have questions or problems! <br>
Contact us during the exercise session, on [ADAM](https://adam.unibas.ch/goto_adam_crs_1266890.html), [Piazza](https://piazza.com/class/kzy15kp8s5t6ku), or [via email](https://sada.dmi.unibas.ch/en/teaching/pids22).

Please submit this exercise sheet on **ADAM**.
Naming conventions:
* Name of notebook: "Exercise sheet 6.ipynb"
* Please put it in a folder structure like this: "your_unibas_short_name/exercise6/Exercise sheet 6.ipynb"
* You can find your unibas short name either...
    * by looking at the filename of the feedback on exercise sheet 2, 3, 4, 5 (from ADAM)
    * or if you didn't submit any exercise sheet yet, log in to viaweb.unibas.ch and you will see your unibas short name after the welcome message
* Then zip this folder and upload it on ADAM.

**Handout date**: 2022/04/11 <br>
**Submission date**: 2022/04/18 <br>
**Discussion date**: 2022/04/19

```python
import pandas as pd
import seaborn as sns
import matplotlib.pyplot as plt
import numpy as np
from nose.tools import assert_is_instance, assert_equal, assert_almost_equal, assert_true
```

## 1 (6 points)
Let's start with some basic probability terms.

### 1a (1 point)
What is $\Omega$ for 3 coin tosses? Use the Python set data structure with tuples to describe each possible outcome. Check the slides for lecture 5 to see what $\Omega$ is.

<!-- grade_id: cell-85f0aa24377f6a04, points: 0, grade: false, solution: true -->
```python
#omega_three_coin_tosses = ...
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-d6087221b91e6736, points: 1, grade: true, solution: false -->
```python
assert_is_instance(omega_three_coin_tosses, set)
assert_true((0,0,0) in omega_three_coin_tosses)
```

### 1b (1 point)
We roll two standard dice.
What is $\Omega$ for "We roll the same number twice?"

<!-- grade_id: cell-e14669da3b7636be, points: 0, grade: false, solution: true -->
```python
#omega_same_number_twice = ...
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-f13015f9c3335967, points: 1, grade: true, solution: false -->
```python
assert_is_instance(omega_same_number_twice, set)
assert_true((1,1) in omega_same_number_twice)
```

### 1c (1 point)
We roll two standard dice. What is the probability for the sum of the two dice being lower than 30?

<!-- grade_id: cell-726e66c3dbc5c1f4, points: 0, grade: false, solution: true -->
```python
#probability_lower_30 = ...
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-9a36a513e0e5666b, points: 1, grade: true, solution: false -->
```python
assert_is_instance(probability_lower_30, float)
```

### 1d (1 point)
We roll two standard dice. What is the probability for the sum of the two dice being lower or equal to 6?

<!-- grade_id: cell-3f1a63a0f6475004, points: 0, grade: false, solution: true -->
```python
#probability_lower_6 = ...
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-912ffb901cab2c66, points: 1, grade: true, solution: false -->
```python
assert_is_instance(probability_lower_6, float)
```

### 1e (2 point)
We roll two standard dice. If the first die shows a "6", what is the probability of the sum of the two dice being equal or higher than 8?

<!-- grade_id: cell-37bbbd825b890acb, points: 0, grade: false, solution: true -->
```python
#probability_higher_8 = ...
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-79bfdfafc1dd07ae, points: 2, grade: true, solution: false -->
```python
assert_is_instance(probability_higher_8, float)
```

## 2 (5 points)
We will work with a dataset about strokes (sourced from [Kaggle](https://www.kaggle.com/datasets/fedesoriano/stroke-prediction-dataset)).

```python
stroke_df = pd.read_csv("healthcare-dataset-stroke-data.csv")
stroke_df
```

### 2a (1 point)
For a patient contained in this dataset, what is the probability that they had a stroke (P(stroke))?

<!-- grade_id: cell-d705987e21fb4422, points: 0, grade: false, solution: true -->
```python
#probability_stroke = ...
# YOUR CODE HERE
raise NotImplementedError()
print(probability_stroke)
```

<!-- grade_id: cell-abfb253fab1b086d, points: 1, grade: true, solution: false -->
```python
assert_is_instance(probability_stroke, float)
```

### 2b (1 point)
For a patient contained in this dataset who had heart disease (heart_disease = 1), what is the probability that they had a stroke (P(stroke|heart_disease))?

<!-- grade_id: cell-5ee38af4cad7f998, points: 0, grade: false, solution: true -->
```python
#probability_stroke_heart_disease = ...
# YOUR CODE HERE
raise NotImplementedError()
print(probability_stroke_heart_disease)
```

<!-- grade_id: cell-4592808d760724cf, points: 1, grade: true, solution: false -->
```python
assert_is_instance(probability_stroke_heart_disease, float)
```

### 2c (3 points)
What is the probability of P(stroke|hypertension)?

Hint: Check the slides for a very similar example about breast cancer

<!-- grade_id: cell-2871bba4629240f8, points: 0, grade: false, solution: true -->
```python
#probability_stroke_hypertension = ...
# YOUR CODE HERE
raise NotImplementedError()
print(probability_stroke_hypertension)
```

<!-- grade_id: cell-1a0d689f67daf28d, points: 3, grade: true, solution: false -->
```python
assert_is_instance(probability_stroke_hypertension, float)
```

### 2d (4 points)
According to the [Mayo Clinic](https://www.mayoclinic.org/diseases-conditions/diabetes/diagnosis-treatment/drc-20371451), a "blood sugar level less than 140 mg/dL (7.8 mmol/L) is normal".

* Make a new column called "high_glucose" that has a 1 if the avg_glucose_level measurement was equal to or higher than 140, and a 0 otherwise

<!-- grade_id: cell-b62c57fe4875d77a, points: 0, grade: false, solution: true -->
```python
# stroke_df["high_glucose"] = ...
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-d2720ca4ba63af38, points: 1, grade: true, solution: false -->
```python
#assert_is_instance(stroke_df["high_glucose"], set)
```

* Use this new column to calculate P(stroke|high_glucose)

<!-- grade_id: cell-96a4ca2cec45e2e3, points: 0, grade: false, solution: true -->
```python
# probability_stroke_high_glucose = ...
# YOUR CODE HERE
raise NotImplementedError()
print(probability_stroke_high_glucose)
```

<!-- grade_id: cell-561eda5d11855e2c, points: 1, grade: true, solution: false -->
```python
assert_is_instance(probability_stroke_high_glucose, float)
```


* Calculate P(high_glucose|stroke)

<!-- grade_id: cell-cc7cf764ecaa7013, points: 0, grade: false, solution: true -->
```python
# probability_high_glucose_stroke = ...
# YOUR CODE HERE
raise NotImplementedError()
print(probability_high_glucose_stroke)
```

<!-- grade_id: cell-7cfc8b9672c5960c, points: 2, grade: true, solution: false -->
```python
assert_is_instance(probability_high_glucose_stroke, float)
```